# Exercise 5: Viscous flow of ice (20 points)

This week's exercise explores the flow of ice through a valley as a channel and down a slope.

### Tips for completing this exercise

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-5-problems-1-3.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

## Problem 1 - Ice flow in an open channel (12 points)

The goal of this problem is to calculate horizontal velocity profiles across a glacier for Newtonian and non-Newtonian fluid flow resulting from a pressure gradient.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating the pressure gradient
- Creating a Python function for the channel flow velocity
- Calculating and plotting ice velocities in a channel
- Including comments that explain what most lines in the code do
- Uploading your notebook and `introqg_functions.py` script file to your GitHub repository for this week's exercise

### Part 0: Copying and testing your script file from Exercise 4 (0 points)

The first task in this problem is to copy your `introqg_functions.py` script file from Exercise 4 to the directory containing this notebook and then run the cell below to ensure it has been copied and is functioning as expected. Note: We will only check some of the functions in this file using the tests below, not all of them.

- Copy your `introqg_functions.py` script file from Exercise 4 to the directory containing this notebook
- Run the tests below

```python
# The test below should work

from nose.tools import assert_equal
from introqg_functions import steady_state_temp, transient_temp

# Test the steady-state temperature calculation function
test_temp1 = steady_state_temp(
    gradient=10.0,
    diffusivity=32.0,
    velocity=0.5,
    depths=10.0)

test_temp2 = steady_state_temp(
    gradient=20.0,
    diffusivity=32.0,
    velocity=0.75,
    depths=15.0)

test_temp3 = transient_temp(
    initial_gradient=20.0,
    diffusivity=32.0,
    velocity=0.5,
    depths=10.0,
    time=5.0)

test_temp4 = transient_temp(
    initial_gradient=15.0,
    diffusivity=32.0,
    velocity=1.5,
    depths=15.0,
    time=15.0)

# Print calculated closure temperature
print(f"Calculated steady-state temperature 1: {test_temp1:7.3f} 째C")
print(f"Calculated steady-state temperature 2: {test_temp2:7.3f} 째C")
print(f"Calculated transient temperature 1: {test_temp3:7.3f} 째C")
print(f"Calculated transient temperature 2: {test_temp4:7.3f} 째C")

# Check that the closure temperature values are correct
assert_equal(round(test_temp1, 3), 92.579)
assert_equal(round(test_temp2, 3), 252.938)
assert_equal(round(test_temp3, 3), 232.717)
assert_equal(round(test_temp4, 3), 465.720)

# Print message if it is safe to continue
print("\nAll tests pass! You are ready to proceed with this exercise.")
```

### Part 1: Background and getting started (0 points)

Before you start you should read through the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-1) for this problem. The other parts of this problem will refer to items described in this theory text.

- Read through the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-1) for this problem

### Part 2: A pressure gradient function? (4 points)

As mentioned in the description of the problem, we will be calculating velocities of ice at the surface as we cross a glacial valley. We can calculate the ice velocity using Equation 10 from the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-1), which is also shown below

\begin{equation}
  \large
  u(y) = \frac{a}{n + 1} \left( \frac{p_{1} - p_{0}}{L} \right)^{n} \left[ \left( \frac{h}{2} \right)^{n + 1} - y^{n + 1} \right]
\end{equation}

where $a$ is a function of the viscosity and temperature (we will ignore temperature for now), $n$ is the power-law exponent, $\frac{p_{1} - p_{0}}{L}$ is the pressure gradient in the channel, $h$ is the channel width, and $y$ is array for distance from the center of the channel. The only problem here is that the value for the pressure gradient $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ is not known. Thus, our first step is to create a function to estimate its value from some of the observed velocities for the Saskatchewan glacier.

- Create a function in your `introqg_functions.py` script file called `channel_pressure` that calculates the value of $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ for a given value of the power-law exponent `n`.
    - You can solve for $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ by rearranging Equation 10 (see above) to solve for $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ and using the condition that the maximum flow velocity $u = 2.41 \times 10^{-6}$ m/s occurs at $y = 0$ (see Table 1 in Part 4).
    - Assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup> and $h = 1400$ m.
    - Your function parameters should include
        - The power-law exponent `exponent`
        - The channel width `width`
        - The viscosity/temperature constant `viscosity_constant`
        - The channel velocity at y = 0 `velocity_max`

<!-- grade_id: problem_2_ag_p2_pressure, points: 0, grade: false, solution: true -->
```python
# Import your channel_pressure function below

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p2_pressure_test1, points: 2, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
pressure1 = channel_pressure(
    exponent=3,
    width=1400.0,
    viscosity_constant=5.0e-24,
    velocity_max=2.41e-6
)

# Print results
print(f"The pressure gradient for exponent=3 is {pressure1:.3f}. Expected value: 8029987.505.")

# Check the other test calculations
assert_equal(round(pressure1, 3), 8029987.505)
```

<!-- grade_id: problem_2_ag_p2_pressure_test2, points: 2, grade: true, solution: false -->
```python
# The visible tests below should work

# Function test calculations
pressure2 = channel_pressure(
    exponent=6,
    width=1400.0,
    viscosity_constant=5.0e-24,
    velocity_max=2.41e-6
)
pressure3 = channel_pressure(
    exponent=1.5,
    width=1400.0,
    viscosity_constant=5.0e-24,
    velocity_max=2.41e-6
)

# Print results
print(f"The pressure gradient for exponent=6 is {pressure2:.3f}. Expected value: 0.041.")
print(f"The pressure gradient for exponent=1.5 is {pressure3:.3f}. Expected value: 92948406117.575.")

# Check the other test calculations
assert_equal(round(pressure2, 3), 0.041)
assert_equal(round(pressure3, 3), 92948406117.575)
```

### Part 3: An ice flow velocity function (4 points)

Your next task is to create a function to calculate the velocity of ice flowing in a channel (Equation 10).

- Create a function in your `introqg_functions.py` script file called `channel_velocity` that calculates the velocity of ice flowing in a channel using Equation 10 from the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-1), which was shown in the previous part.
    - Your function should use the `channel_pressure` function created in the previous part.
    - As in the previous part you should assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>, $h = 1400$ m, and that the flow is symmetric about $y = 0$.
    - Your function parameters should include
        - The power-law exponent `exponent`
        - The channel width array `y`
        - The channel width `width`
        - The viscosity/temperature constant `viscosity_constant`
        - The channel velocity at y = 0 `velocity_max`

<!-- grade_id: problem_2_ag_p3_velocity, points: 0, grade: false, solution: true -->
```python
# Import your channel_velocity function below

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p3_velocity_test1, points: 2, grade: true, solution: false -->
```python
# The visible tests below should work
import numpy as np
from nose.tools import assert_equal

# Fake data
y1 = np.array([-500.0, -250.0, 0.0, 200.0, 400.0, 600.0])
velocity1 = channel_velocity(exponent=2, y=y1)

# Print results
print(f"The channel velocity at y1[1] for exponent=2 is {velocity1[1]:.5e} m/s. Expected value: 2.30022e-06 m/s.")

# Check the other test calculations
assert_equal(round(velocity1[-1], 8), 8.9e-07)
assert_equal(round(velocity1[3], 8), 2.35e-06)
```

<!-- grade_id: problem_2_ag_p3_velocity_test2, points: 2, grade: true, solution: false -->
```python
# The visible tests below should work

# Fake data
y2 = np.array([-300, -150.0, 0.0, 100.0, 200.0, 300.0])
velocity2 = channel_velocity(exponent=4, y=y2)

# Print results
print(f"The channel velocity at y2[1] for exponent=4 is {velocity2[1]:.5e} m/s. Expected value: 2.40891e-06 m/s.")

# Check the other test calculations
assert_equal(round(velocity2[-1], 8), 2.38e-06)
assert_equal(round(velocity2[3], 8), 2.41e-06)
```

### Part 4: Calculating ice flow velocities (2 points)

Your next task is to calculate the flow velocity for ice flowing in a channel using your `channel_velocity` function.

- Create a new pandas DataFrame called `channel` to store channel velocities for 1400-m-wide ice channel
    - Use `y` as the index for your `channel` DataFrame
    - **Hint**: You should define $width = 1400$ m and assume that the flow is symmetric about $y = 0$ and define `y` with a length of 201
- Calculate the velocity of ice in a channel as a non-Newtonian fluid with power-law exponents of $exponent$ = 2, 3, 4, and 5.
    - As above, you can assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>
- Store the output from each call of the `channel_velocity` function for each value of $exponent$ in your `channel` DataFrame
    - You should label the new column in your DataFrame using the power-law exponent. For example, you could have a column called `exponent = 2` for a power-law exponent of 2.
    - **Note**: You should convert your channel velocities from m/s to m/yr before storing them in the `channel` DataFrame!
    - The hints for this week's exercise may be helpful for this task

<!-- grade_id: problem_2_ag_p4_velocities, points: 0, grade: false, solution: true -->
```python
# Import pandas
import pandas as pd

# Enter your values below
width = None
y = None
channel = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p4_velocities_test1, points: 2, grade: true, solution: false -->
```python
# These tests should work
print(f"The channel shape is {channel.shape}. Expected value: (201, 4).")
print(f"The velocity at channel.iloc[51,0] is {channel.iloc[51,0]:.4f} m/yr. Expected value: 67.1062 m/yr.")

```

### Part 5: Plotting your results (2 points)

Next you can create a plot of the velocity calculations you've done.

- Create a plot of the velocity of the ice for all 4 power-law exponents
    - Be sure to label your axes and add a title.
    - Also be sure your line legend makes it clear which profile is which
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_2_mg_p5_plot, points: 1.5, grade: true, solution: true -->
```python
# Add your plot below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p5_plot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 2 - Comparing our model to the Saskatchewan glacier velocities (8 points)

In this problem we will use the `channel_velocity` function from Problem 2 to compare predicted ice velocities to glacier flow velocities observed from the [Saskatchewan glacier near Banff in Alberta, Canada](https://goo.gl/maps/R9S48J4KYPr) (Figure 1). The Saskatchewan glacier is 1400 m wide and part of a large ice field known as the Columbia Icefield.

**Your score on this problem will be based on following criteria**:

- Reading in the data file
- Comparing predictions for the ice flow velocity to observations from the data file
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook and `introqg_functions.py` script file to your GitHub repository for this week's exercise

![Saskatchewan glacier](https://upload.wikimedia.org/wikipedia/commons/a/a5/Saskatchewan_Glacier.jpg)<br/>
*Figure 1. Saskatchewan glacier.*.

### Part 1: Reading the data file (2 points)

Distance (m) | Velocity (m/yr) | Velocity (m/s)
------------ | -------------- | --------------
-660 | 12 | 3.80E-07
-640 | 22 | 7.00E-07
-570 | 42 | 1.33E-06
-460 | 63 | 2.00E-06
-220 | 74 | 2.35E-06
40 | 76 | 2.41E-06
180 | 74 | 2.35E-06
260 | 72 | 2.28E-06
500 | 51 | 1.62E-06

*Table 1. Velocity measurements across the Saskatchewan glacier.*

The file [`data/sask_glacier_velo.txt`](data/sask_glacier_velo.txt) contains a series of surface velocities measured at various locations across the glacier (Table 1). Your task now is to load in this velocity data.

- Read the data file [`data/sask_glacier_velo.txt`](data/sask_glacier_velo.txt) using pandas into a DataFrame called `sask`
    - Set the column `'Distance (m)'` as the index column

<!-- grade_id: problem_3_ag_p2_read, points: 0, grade: false, solution: true -->
```python
# Load the data file
sask = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_ag_p2_read_test1, points: 2, grade: true, solution: false -->
```python
# This test should work
print(f"The first 5 rows of the data file are:\n\n{sask.head()}")

```

### Part 2: Comparing to observations (4 points)

The next task is to calculate ice velocities at locations corresponding to the observed velocity measurements and compare the predictions to the data in the data file.

- Loop over the power-law exponent values $exponent=2$ to $exponent=5$ and calculate the velocity of ice using the `channel_velocity` function at the positions where the velocity of the Saskatchewan glacier has been observed in the data file
    - Use the default function values other than the power-law exponent and values for `y`
    - You should label the new column in your `sask` DataFrame using the power-law exponent. For example, you could have a column called `exponent = 2` for a power-law exponent of 2.
    - **Note**: You should convert your channel velocities from m/s to m/yr before storing them in the `sask` DataFrame!
- For each set of predicted ice flow velocities, calculate the goodness-of-fit to the observed ice velocities and append this to the `gof_sask` list
    - We will be using your `chi_squared` function from back in Exercise 2
    - We are assuming an arbitrary standard deviation for the observed velocities of 2 m/yr.
        - When you call your `chi_squared` function you can pass the variable `sask_stdev` for the standard deviations. It is defined in the cell below.

<!-- grade_id: problem_3_ag_p3_gof, points: 0, grade: false, solution: true -->
```python
# Import items we need
import numpy as np
from introqg_functions import chi_squared

# Create array of standard deviations
sask_stdev = np.ones(len(sask)) * 2.0

# Create empty list for goodness-of-fit values
gof_sask = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_ag_p3_gof_test1, points: 2, grade: true, solution: false -->
```python
# This test should work
print(f"The first calculated ice velocity is {sask.iloc[0,2]:.2f} m/yr. Expected value: 12.31 m/yr.")

```

<!-- grade_id: problem_3_ag_p3_gof_test2, points: 2, grade: true, solution: false -->
```python
# This test should work
print(f"The first goodness of fit value is {gof_sask[0]:.2f}. Expected value: 4.06.")

```

### Part 3: Plotting your results (2 points)

The last task in this problem is to plot the predicted ice flow velocities along with the observed values and show the best goodness of fit as text on the plot.

- Plot the measured velocities along with the 4 predicted velocity profiles you have calculated **in meters per year (m/yr)**
    - For the predicted velocities, use the values from the `channel` DataFrame from Problem 1
    - For the observed velocities, use the values from the `sask` DataFrame
        - Plot these as black circles
    - Be sure to label your axes and include a title
    - Also add the best goodness-of-fit value as text on the plot along with the corresponding power-law exponent value
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_3_mg_p4_plot, points: 1.5, grade: true, solution: true -->
```python
# Plot your results below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_mg_p4_plot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 3 - Non-Newtonian ice flow down an inclined plane (*Optional*, 0 points)

The [Athabasca glacier](https://goo.gl/maps/HggYfoKxEUQ2) (Figure 2) is another glacier in the Columbia Icefield, which will be the focus in this problem.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating the gravitational force
- Creating a Python function for calculating the ice velocity
- Reading in the data file
- Comparing predictions for the ice flow velocity to observations from the data file
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook and `introqg_functions.py` script file to your GitHub repository for this week's exercise

![Athabasca glacier](https://upload.wikimedia.org/wikipedia/commons/4/41/Icefields_parkway.jpg)<br/>
*Figure 2. Athabasca glacier.*

### Part 1: Background and theory (0 points)

First read through the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-2) for this problem before proceeding.

- Read through the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-2) for this problem

### Part 2: A gravitational force function? (0 points)

In this problem we will be calculating the flow of ice down a slope resulting from a gravitational forces on the glacier. As above, we will consider both Newtonian and non-Newtonian fluids, this time using Equation 19 from the [background and theory for this problem](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-2), which is also shown below

\begin{equation}
  \large
  u(z) = u_{\mathrm{b}} + \frac{a}{n + 1} \gamma_{x}^{n} \left[ h^{n + 1} - (h - z)^{n + 1} \right]
\end{equation}

where $u_{\mathrm{b}}$ is the basal sliding velocity of the glacier, $a$ is a function of the viscosity and temperature (we will ignore temperature for now), $n$ is the power-law exponent, $\gamma_{x}$ is the downslope component of the gravitational force, $h$ is the thickness of the ice perpendicular to the slope, and $z$ is the array of distances from the base of the glacier. Similar to Problem 3, the value for the gravitational force $\gamma_{x}$ is not known, but we can create a function to estimate its value.

- Create a function in your `introqg_functions.py` script file called `gravity_force` that calculates the value of $\gamma_{x}^{n}$ for a given value of the power-law exponent `n`.
    - You can solve for $\gamma_{x}^{n}$ by rearranging Equation 19 (see above) to solve for $\gamma_{x}^{n}$ and using the condition that the maximum flow velocity $u = 9.1 \times 10^{-7}$ m/s occurs at $z = h = 209$ m (see Table 2 in Part 4).
    - Note that you should assume the basal sliding velocity $u_{\mathrm{b}}$ is zero for the Newtonian case and equal to the observed sliding velocity $u_{\mathrm{b}} = 1.30 \times 10^{7}$ m/s for the non-Newtonian cases (see Table 2 in Part 4).
    - Again, assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup> and $h = 209$ m.
    - Your function parameters should include
        - The power-law exponent `exponent`
        - The ice thickness `thickness`
        - The viscosity/temperature constant `viscosity_constant`
        - The maximum ice flow velocity `velocity_max`
        - The basal sliding velocity `velocity_basal`

<!-- grade_id: problem_4_ag_p2_gamma_x, points: 0, grade: false, solution: true -->
```python
# Import your gravity_force function below

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p2_gamma_x_test1, points: 0, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
force1 = gravity_force(
    exponent=1,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)

# Print results
print(f"The gravity force for n=1 is {force1:.2f}. Expected value: 8333142556260.16.")

# Check the other test calculations
assert_equal(round(force1, 2), 8333142556260.16)
```

<!-- grade_id: problem_4_ag_p2_gamma_x_test2, points: 0, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
force2 = gravity_force(
    exponent=5,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)
force3 = gravity_force(
    exponent=7,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)

# Print results
print(f"The gravity force for n=5 is {force2:.2f}. Expected value: 11230.48.")

# Check the other test calculations
assert_equal(round(force2, 2), 11230.48)
assert_equal(round(force3, 2), 0.34)
```

### Part 3: An ice flow velocity function, again (0 points)

Your next task is to create a function to calculate the velocity profile for ice flowing down a slope (Equation 19), shown in Part 2.

- Create a function in your `introqg_functions.py` script file called `slope_velocity` that calculates the velocity of ice flowing down a slope using Equation 19 from the [background and theory](https://introqg-site.readthedocs.io/en/latest/lessons/L5/exercise-5-theory.html#problem-2), which was shown in the previous part.
    - Your function should use the `gravity_force` function created in the previous part.
    - As in the previous part you should assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>, $h = 209$ m, and the maximum flow velocity $u_{\mathrm{max}} = 9.1 \times 10^{-7}$ m/s.
    - Also you should again assume the basal sliding velocity $u_{\mathrm{b}}$ is zero for the Newtonian case and equal to the observed sliding velocity $u_{\mathrm{b}} = 1.30 \times 10^{7}$ m/s for the non-Newtonian cases (see Table 2 in Part 4).
    - Your function parameters should include
        - The power-law exponent `exponent`
        - The channel thickness array `z`
        - The ice thickness `thickness`
        - The viscosity/temperature constant `viscosity_constant`
        - The maximum ice flow velocity `velocity_max`
        - The basal sliding velocity `velocity_basal`

<!-- grade_id: problem_4_ag_p3_slope_velocity, points: 0, grade: false, solution: true -->
```python
# Import your slope_velocity function below

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p3_slope_velocity_test1, points: 0, grade: true, solution: false -->
```python
# The visible tests below should work
import numpy as np
from nose.tools import assert_equal

# Function test calculations
z1 = np.array([0.0, 21.0, 74.0, 104.0, 159.0, 209.0])
velocity1 = slope_velocity(
    exponent=3,
    z=z1,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)

# Print results
print(f"The velocity velocity1[4] for exponent=3 is {velocity1[4]:.8f} m/s. Expected value: 0.00000091 m/s.")

# Check the other test calculations
assert_equal(round(velocity1[-1], 8), 9.1e-07)
assert_equal(round(velocity1[3], 8), 8.6e-07)
```

<!-- grade_id: problem_4_ag_p3_slope_velocity_test2, points: 0, grade: true, solution: false -->
```python
# The visible tests below should work

# Function test calculations
velocity2 = slope_velocity(
    exponent=1,
    z=z1,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)
velocity3 = slope_velocity(
    exponent=7,
    z=z1,
    thickness=209.0,
    viscosity_constant=5.0e-24,
    velocity_max=9.1e-7,
    velocity_basal=1.3e-7
)

# Print results
print(f"The velocity velocity2[4] for exponent=1 is {velocity2[4]:.8f} m/s. Expected value: 0.00000086 m/s.")

# Check the other test calculations
assert_equal(round(velocity2[-1], 8), 9.1e-07)
assert_equal(round(velocity2[3], 8), 6.8e-07)
assert_equal(round(velocity3[-1], 9), 9.1e-07)
assert_equal(round(velocity3[3], 9), 9.07e-07)
```

### Part 4: Reading the data file (0 points)

Depth from surface (m) | Height from base (m) | Velocity (m/yr) | Velocity (m/s)
---------------------- | -------------------- | -------------- | --------------
0 | 209 | 28.6 | 9.10E-07
15 | 195 | 28.5 | 9.00E-07
30 | 180 | 28.5 | 9.00E-07
45 | 165 | 28.4 | 9.00E-07
60 | 150 | 28.2 | 8.90E-07
75 | 135 | 28.0 | 8.90E-07
90 | 120 | 27.7 | 8.80E-07
105 | 105 | 27.2 | 8.60E-07
120 | 90 | 26.5 | 8.40E-07
135 | 75 | 25.5 | 8.10E-07
150 | 60 | 24.0 | 7.60E-07
165 | 45 | 21.5 | 6.80E-07
180 | 30 | 17.5 | 5.50E-07
195 | 15 | 10 | 3.20E-07
209 | 0 | 4 | 1.30E-07

*Table 2. Velocities measured from a vertical profile through Athabasca glacier.*

A vertical velocity profile for the Athabasca glacier has been measured and the measurements are in the file [`data/atha_glacier_velo.txt`](data/atha_glacier_velo.txt) (Table 2). Your task now is to load in this velocity data.

- Read the data file [`data/atha_glacier_velo.txt`](data/sask_glacier_velo.txt) using pandas into a DataFrame called `athabasca`
    - **Note**: Do NOT assign an index column, this will cause trouble when plotting the data

<!-- grade_id: problem_4_ag_p4_data, points: 0, grade: false, solution: true -->
```python
# Import pandas
import pandas as pd

# Read the data file below
athabasca = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p4_data_test1, points: 0, grade: true, solution: false -->
```python
# This test should work
print(f"The first 5 rows of the data file are\n\n{athabasca.head()}")

```

### Part 5: Comparing to observations, again (0 points)

The next task is to calculate ice velocities at locations corresponding to the observed velocity measurements and compare the predictions to the data in the data file.

- Loop over the power-law exponent values $n=1$ to $n=5$ and calculate the velocity of ice using the `slope_velocity` function at the positions where the velocity of the Athabasca glacier has been observed in the data file
    - Use the default function values other than the power-law exponent and values for `z`
        - **Hint**: Use the column `'Height from base (m)'` for the values for `z`
    - You should label the new column in your `athabasca` DataFrame using the power-law exponent. For example, you could have a column called `n = 2` for a power-law exponent of 2.
    - **Note**: You should convert your ice velocities from m/s to m/yr before storing them in the `athabasca` DataFrame!
- For each set of predicted ice flow velocities, calculate the goodness-of-fit to the observed ice velocities and append this to the `gof_athabasca` list
    - We will be using your `chi_squared` function from back in Exercise 2
    - We are assuming an arbitrary standard deviation for the observed velocities of 1 m/yr.
        - When you call your `chi_squared` function you can pass the variable `athabasca_stdev` for the standard deviations. It is defined in the cell below.

<!-- grade_id: problem_4_ag_p5_gof, points: 0, grade: false, solution: true -->
```python
# Import items we need
import numpy as np
from introqg_functions import chi_squared

# Create array of standard deviations
athabasca_stdev = np.ones(len(athabasca))

# Create empty list for goodness-of-fit values
gof_athabasca = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p5_gof_test1, points: 0, grade: true, solution: false -->
```python
# This test should work
print(f"The first calculated ice velocity is {athabasca.iloc[0,4]:.2f} m/yr. Expected value: 28.72 m/yr.")

```

<!-- grade_id: problem_4_ag_p5_gof_test2, points: 0, grade: true, solution: false -->
```python
# This test should work
print(f"The last goodness of fit is {gof_athabasca[-1]:.2f}. Expected value: 1.18.")

```

### Part 6: Plotting your results (0 points)

The last task in this problem is to again plot the predicted ice flow velocities along with the observed values and show the best goodness of fit as text on the plot.

- Plot the measured velocities along with the 5 predicted velocity profiles you have calculated **in meters per year (m/yr)** on the x-axis as a function of their **height from the base** on the y-axis
    - For the predicted velocities, use the values from the `athabasca` DataFrame
    - For the observed velocities, also use the values from the `athabasca` DataFrame
        - Plot these as black circles connected by a black line
    - Be sure to label your axes and include a title
    - Also add the best goodness-of-fit value as text on the plot along with the corresponding power-law exponent value
    - **Hint**: You should use the `label` parameter for your plots to provide the text to use for the legend
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_4_mg_p6_plot, points: 0, grade: true, solution: true -->
```python
# Plot your results below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_mg_p6_plot_caption, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Optional reflection questions (0 points)

1. Based on your models, what is the most likely power-law exponent for glacial ice?
2. Looking at your plots, how sensitive is the velocity of ice to the power-law exponent?

<!-- grade_id: summary_questions_mg, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE